Transdiagnostic compulsivity is associated with reduced reminder setting, only partially attributable to overconfidence

Author response:

The following is the authors’ response to the current reviews.

Reviewer #1:

(1) To improve the clarity of the work, I suggest a final note to the authors to say more explicitly that objective accuracy has a finer resolution *due to the number of "special circles" per trial* in their task. This task detail got lost in my read of the manuscript, and confused me with respect to the resolution of each accuracy measure.

We agree with the reviewer that this would be a useful clarification and have therefore added the following statement to the Methods section on p. 20:

“It should be noted that the OIP has a slightly finer resolution due to the number of special circles per trial.”

(2) Similarly for clarification, they could point out that their exclusion criteria removes subjects that have lower OIP than their AIP analysis allows (which is good for comparison between OIP and AIP). Thus, it removes the possibility that very poor performing subjects (OIP) are forced to have a higher than actual AIP due to the range).

We agree this would be a useful statement to add and have included the following sentence in the Supplement on p. 8:

“Such a restriction of the threshold parameter was intended to increase the comparability between AIP and OIP, and hence improved the calculation of the reminder bias.”

The following is the authors’ response to the previous reviews.

Reviewer #1:

(1) Upon reading their response to the question I had regarding AIP and OIP, a few more questions came up regarding OIP, AIP, how they're calculations differ, and how the latter was computed in R. I hope these help readers to clarify how to interpret these key measures, and the hypotheses that rely upon them.

Regarding fitting, and in relation to power, is16 queries adequate to estimate an AIP using the R's quickpsy? That is, assuming some noise in the choice process, how recoverable is a true indifference points from 16 trials? If there's a parameter recovery analysis (ie generating choice via the fitting parameters, which will have built-in stochasticity, and seeing how well you recover the parameter) of interest would be helpful. It may help to characterize why the present study might differ from prior studies (maybe a power issue here).

The reviewer is absolutely correct that we should have provided more detail when describing our fitting procedure for the psychometric curves. We have now addressed this by adding the following statements to the Methods section and Supplement:

Page 20 in the main manuscript: “Fitting was done using the quickpsy package in R and more detail is given in the Supplement.”

Pages 8 and 9 in the Supplement:

“Psychometric curve fitting

We used the quickpsy package in R to fit psychometric curves to each participant’s choice data to derive their actual indifference point (AIP), which was operationalised as the threshold parameter when predicting reminder choices from target values. We restricted the possible parameter ranges from 2 to 9 for the threshold parameter and from 1 to 500 for the slope parameter, based on the task’s properties and pilot data. Apart from those parameter ranges, we used only default settings of the quickpsy() function.

Each participant has only 16 trials (2 for each target value) contribute to the curve fitting. To understand the robustness of the AIP based on such limited data, we conducted a parameter recovery analysis. We simulated 16 trials based on each psychometric function and re-ran the curve fitting based on those simulated choices. There was close correspondence between the actual and recovered threshold parameters (or AIPs) with a correlation of r = 0.97, p < 0.001 (see also Figure S1). In contrast, the slope parameter—which was not central to any of our analyses—exhibited greater variability during the initial fitting. This increased uncertainty likely contributed to its poor recovery in the simulation, as evidenced by a near-zero correlation (r = −0.01, p = 0.82).”

(2) Along these lines, it would be helpful for the reader to actually see the individual psychometric curve, now how quickpsy was used (did you fit left and right asymptotes), etc, to understand how that fitting procedure works and how the assumptions of the fitting procedure compare to what can be gleaned through seeing the choice curves plotted.

As stated above, we used default settings of the quickpsy() function and hence assumed symmetric asymptotes at 0 and 1. However, the reviewer mentions “left and right asymptotes”, so maybe this question is about restricting the possible parameter range for the threshold, which we restricted to values from 2 to 9, as described above.

Regarding the individual curves, we have now include the following statement on page 9 in the Supplement: “Figures S2 to S31 show the individual psychometric curves that were estimated for each participant.” Please refer to the Supplement for the added figures.

(3) A more full explanation of quickpsy, its parameters, and how choice curves look might also generate interesting further questions to think about with respect to biases and compulsivity. Two individuals might have similar indifference points, but an asymptote might reflect a bias to always have some percent chance of for example to take the reminders even at the lowest offer available for them.

We agree that this is an interesting focus which we will keep in mind for future studies.

(4) Regarding comparing OIP to AIP:

For OIP, as far as I can understand, the resolution of it is decreased compared to AIP. Accuracies for OIP can only be 0/4,1/4,2/4,3/4, or 4/4. Yet, the resolution for AIP is the full range of offers (2 to 9) with respect to the parameter of interest (the indifference point). Could this bias the estimation of OIP (for instance, someone who scored 25% might actually be much closer to either 50 or 0, but we can't tell due to resolution?

As mentioned in response to comment (1), we restricted the parameter range for the thresholds to 2 to 9 to increase comparability. The reviewer is right to point out that the OIP still has lower resolution than the AIP, which is one of the downsides of having a shortened paradigm (cf. the longer version in Gilbert et al., 2019), which is optimised for online testing, especially if used in combination with additional questionnaires. We have no reason to believe though that this could have led to any bias, especially none that would contribute to the individual differences which are the main focus of our study.

Gilbert, S. J., Bird, A., Carpenter, J. M., Fleming, S. M., Sachdeva, C., & Tsai, P.-C. (2020). Optimal use of reminders: Metacognition, effort, and cognitive offloading. Journal of Experimental Psychology: General, 149(3), 501–517. https://doi.org/10.1037/xge0000652

(5) Additionally, it seems like the upper and lower bounds of OIP (0 and 10) differ from AIP (2 and 9). Could this also introduce bias (for example, if someone terrible performance, the mean would artificially be higher under AIP than OIP because the smallest indifference point is 2 under AIP, but could be 0 under OIP.

See our response to comment (1), we fixed the range to 2 to 9 (which was the range of target values used in our study).

(6) Finally seeing how CIT actually corresponds to accuracy overall (not a relative measure like AIP compared to OIP) I think would also be helpful as this is related to most points noted above.

We included the suggested test as an exploratory analysis on pages 42-43 in the Supplement: “Third, we were interested in how the transdiagnostic phenotypes would correspond to performance. We therefore fitted a model which predicted internal accuracy (that is, unaided task performance on trials where no reminders could be used) from AD, CIT, and the other covariates (age, education and gender). We found that neither AD, β = -0.02, SE = 0.05, t = 0.44, p = 0.658, nor CIT, β = -0.03, SE = 0.05, t = -0.66, p = 0.510, predicted internal accuracy.

The full results can be found in Table S13 as well as in Figure S32.”

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: In the current study, we explored the behavioural and cognitive correlates of the transdiagnostic trait ‘compulsive behaviour and intrusive thought’ (CIT). CIT is associated with impaired metacognition, which in turn has been associated with cognitive offloading behaviours such as external reminder-setting that play a key role in fulfilling cognitive goals. In an online study (N=600) we investigated individual differences in compulsivity, metacognition, and external reminder-usage. Compulsive individuals had reduced preference for external reminders. This was partially, but not fully, attributable to their relative overconfidence. In contrast to previous studies, we found no evidence for an impaired confidence-action link: compulsive individuals used their metacognition to guide offloading just as much as their noncompulsive counterparts. Given the compensatory nature of cognitive offloading, our findings imply that compulsive individuals are at increased risk of inadequate external memory support. Along with transdiagnostic variation in the general population, this finding could also have implications for clinical conditions, such as OCD.

Revised version: In the current study, we explored the behavioural and cognitive correlates of the transdiagnostic trait ‘compulsive behaviour and intrusive thought’ (CIT). CIT is associated with impaired metacognition, which in turn has been associated with cognitive offloading behaviours such as external reminder-setting that play a key role in fulfilling cognitive goals. In an online study (N=600) we investigated individual differences in compulsivity, metacognition, and external reminder-usage. Compulsive individuals had reduced preference for external reminders. This was partially, but not fully, attributable to their relative overconfidence. In contrast to previous studies, we found no evidence for an impaired confidence-action link: compulsive individuals used their metacognition to guide offloading just as much as their non-compulsive counterparts. Given the compensatory nature of cognitive offloading, our findings imply that compulsive individuals are at increased risk of inadequate external memory support. Along with transdiagnostic variation in the general population, this finding could also have implications for clinical conditions, such as OCD.

Introduction

In recent studies of clinically relevant individual differences, there has been a paradigm shift towards of the study of transdiagnostic traits, challenging the traditional, diagnostic approach. Using factor analysis, temporally stable (see Fox et al., 2023a; Sookud, Martin, Gillan, & Wise, 2024), transdiagnostic phenotypes can be extracted from extensive symptom datasets (Wise, Robinson, & Gillan, 2023). These traits are not confined to a single clinical diagnosis but instead can span a range of conditions, at the same time addressing the diagnostic heterogeneity within conditions, such as obsessive-compulsive disorder (OCD; e.g., Gillan et al., 2016; Wise & Dolan, 2020). There are obvious practical benefits of these methodologies, such as their potential to reduce the clinical burden by making the treatment of comorbid conditions more efficient and effective (Harvey, Watkins, Mansell, & Shafran, 2004). At the same time, they contribute valuable insights into mental health conditions by increasing statistical power and opening new avenues of inquiry (Dalgleish, Black, Johnston, & Bevan, 2020).

In the present study, our focus lies on the latter with the goal to investigate the downstream cognitive and behavioural correlates associated with transdiagnostic compulsivity. This symptom dimension represents a clinical concept characterized by an inability to regulate repetitive behaviours that are harmful to oneself, commonly observed in a variety of conditions, particularly obsessive-compulsive disorder (OCD), schizophrenia, addiction, and eating disorders. Previous research links transdiagnostic compulsivity to impairments in metacognition, defined as thinking about one’s own thoughts, encompassing a broad spectrum of self-reflective signals, such as feelings of confidence (e.g., Rouault, Seow, Gillan & Fleming, 2018; Seow & Gillan, 2020; Benwell, Mohr, Wallberg, Kouadio, & Ince, 2022; Fox et al., 2023a; Fox et al., 2023b; Hoven, Luigjes, Denys, Rouault, van Holst, 2023a). Other studies have shown that metacognitive signals such as feelings of confidence guide cognitive offloading strategies like setting external reminders as memory aids (e.g., Gilbert, 2015; Boldt & Gilbert, 2019). Here, we aim to bridge these two literatures by investigating compulsivity, metacognition, and cognitive offloading within a single experimental paradigm. While compulsivity and cognitive offloading have both separately been linked to metacognition, the relationship between the two - with metacognition as a potential mediating factor - has not previously been systematically examined. This matters because cognitive offloading plays an integral role in our daily lives and is a key contributor to our effectiveness as cognitive agents (Gilbert et al., 2023).

Metacognition guides reminder setting

Reminders constitute an example of cognitive offloading, defined as the use of physical action to reduce the cognitive demands of a task. By offloading memory demands this way, we not only increase the likelihood of successfully completing tasks (Boldt & Gilbert, 2019), but we may also free up cognitive resources for other activities (Dupont, Zhu & Gilbert, 2022). Choosing between setting a reminder and relying on memory is not a trivial matter. Prior research has emphasized the role of metacognition in determining when individuals resort to cognitive offloading (Gilbert et al., 2023; Gilbert, 2015; Boldt & Gilbert, 2019; Sachdeva & Gilbert, 2020; Risko & Gilbert, 2016): People tend to set more reminders when they feel less confident. In other words, people tend to set reminders when they think that they will forget, and this effect holds even after taking into account actual memory ability (e.g., Boldt & Gilbert, 2019). The link between confidence and offloading is observed both for situational fluctuations in confidence due to varying task difficulties (state variable; Boldt & Gilbert, 2022), and for a general predisposition towards over- or underconfidence (trait variable; Boldt & Gilbert, 2019).

Metacognition, compulsivity, and checking behaviours

Given the known metacognitive impairments associated with compulsivity, changes in reminder-setting behaviour are plausible. More specifically, individuals characterized by transdiagnostic compulsivity have been consistently found to exhibit overconfidence (Rouault, Seow, Gillan & Fleming, 2018; Seow & Gillan, 2020; Benwell, Mohr, Wallberg, Kouadio, & Ince, 2022; Fox et al., 2023a; Fox et al., 2023b; Hoven et al., 2023a). If we consider the link between reminder setting and confidence, this implies a reduced likelihood of utilizing external aids, such as reminders. However, while transdiagnostic compulsivity is liked to overconfidence, the opposite pattern of underconfidence is more common in patients with OCD, a compulsive disorder (as reviewed in Hoven, Lebreton, Engelmann, Denys, Luigjes, & van Holst, 2019). Recent research suggests that metacognitive impairments in transdiagnostic compulsivity and OCD may originate from different mechanisms (Hoven, Rouault, von Holst, & Luigjes, 2023c), advising caution against broad generalizations between these groups. It should also be noted that the composite measure of transdiagnostic compulsivity includes questionnaire items linked not only with OCD but also other clinical conditions such as eating disorders (Tasca et al., 2011; Gillan et al., 2016). This results in an overlap between transdiagnostic compulsivity and other traits such as rigid perfectionism.

Despite opposite trends in metacognitive monitoring performance (under-versus overconfidence), individuals high in transdiagnostic compulsivity and those with a diagnosis of OCD show similar impairments in metacognitive control, characterized by a disrupted connection between confidence and future actions (Seow & Gillan, 2020; Vaghi, Luyckx, Sule, Fineberg, Robbins, & De Martino, 2017). Metacognitive impairments are also central to explanations of compulsive behaviours, notably in OCD patients. In such patients, compulsivity can manifest in the form of checking behaviours, for example, checking that doors are locked or that appliances are switched off (Den Ouden et al., 2022). Whilst checking behaviours are also present in other compulsive disorders (e.g., ‘body checking’ in eating disorders; Mountford, Haase, & Waller, 2006), in OCD, these checks are often repetitive and ritualised and are typically associated with obsessive thoughts. However, the exact function that checking compulsions serve is unclear; patients commonly report that they have the aim of reducing anxiety generally, preventing a feared consequence from taking place or that they are performed automatically and without thinking (Starcevic et al., 2011). Understanding these motivators has been challenging as studies rely on self-report of often highly individual real-world behaviours.

Some research argues that OCD patients’ checking arises from low memory confidence despite intact memory (Tolin, Abramowitz, Brigidi, Amir, Street, & Foa, 2001). Our study has the potential to shed some light on the link between confidence and checking: While checking behaviours can be seen as a way of ensuring that a necessary action was performed in the past, reminder-setting is a way of ensuring that a necessary action will be performed in the future. In other words, a reminder can serve as a future checkpoint that allows us to revisit a task at an appropriate time to complete it, perhaps by setting an alarm on our phone, jotting down a note, or strategically placing a related object somewhere visible. Given these insights, one might expect an increased reliance on reminders among OCD patients as they strive to establish more checkpoints. By contrast, seeing as transdiagnostic compulsivity is associated with increased confidence, this could be associated with the opposite pattern: a decreased reliance on reminders.

Three possible mechanisms for changes in reminder setting

If, as hypothesised, compulsivity is linked with altered reminder-setting, this could be attributed to at least three underlying mechanisms. First is the Metacognitive Control Mechanism: Previous research has found that more compulsive individuals tend to have impaired metacognitive control (Seow & Gillan, 2020), meaning they use metacognitive signals to a lesser extent to guide future behaviour. Compulsivity is a hallmark symptom of OCD and similar deficits in metacognitive control have been observed in a case-control studies comparing OCD patients with healthy controls examining how confidence and action are correlated (Vaghi et al., 2017; though see also Hoven, Mulder, Denys, van Holst & Luigjes, 2023b; Marzuki et al., 2022). In the context of our study, a Metacognitive Control Mechanism would be reflected in a disrupted relationship between confidence levels and their tendency to set reminders (i.e., the interaction between the bias to be over- or underconfident and transdiagnostic CIT in a regression model predicting a bias to set reminders).

Second, more compulsive individuals might conceivably differ in their reminder-setting strategies due to an altered level of confidence. We call this the Metacognitive Monitoring Mechanism, which suggests that the issue arises when forming the confidence signal, rather than in its behavioural application (for clarification on metacognitive monitoring versus control in cognitive offloading, see Boldt & Gilbert, 2022). Prior evidence exists for overconfidence in compulsivity (Rouault et al., 2018; Seow & Gillan, 2020; Benwell et al., 2022; Fox et al., 2023a; Fox et al., 2023b; Hoven et al., 2023a), which would therefore result in fewer reminders.

Lastly, there could be a direct link between compulsivity and reminder-usage, independent of any metacognitive influence. We refer to this as the Direct Mechanism and it constitutes any possible influences that affect reminder setting in highly-compulsive CIT participants outside of metacognitive mechanisms, such as perfectionism and the wish to control the task without external aids. Our study aims to differentiate between these three mechanisms. Back when we preregistered our hypotheses only a limited number of studies about confidence and transdiagnostic CIT were available. This resulted in us hypothesising to find support for the Metacognitive Control Mechanism and that highly compulsive individuals would offload more due to an increased need for checkpoints. Both of these hypotheses turned out to be incorrect.

Anxious-depressed transdiagnostic phenotype

As well as investigating individual differences in compulsivity, we also measured an anxious-depression factor. Based on previous findings, we predicted opposite influence of these two factors on confidence. Whereas compulsivity has been linked to increased confidence, anxious-depressed individuals typically display relative underconfidence (Rouault, Seow, Gillan & Fleming, 2018; Seow & Gillan, 2020; Benwell, Mohr, Wallberg, Kouadio, & Ince, 2022). By taking a transdiagnostic approach, we were able to jointly investigate the influence of these two factors of confidence which could potentially cancel out if they were investigated separately.

Online reminder-setting task

In the present preregistered study, we asked 600 participants drawn from the general population to complete several individual differences questionnaires. These responses were then weighted to produce both a ‘compulsive behaviour and intrusive thought’ (CIT) factor and an ‘anxious-depression’ (AD) factor (Gillan, et al., 2016; Wise & Dolan, 2020). Participants’ scores on these factors were then correlated with their behaviour in a reminder-setting task, which was a modified, 20-minute version of the online reminder setting task developed by Gilbert and colleagues (2020; Figure 6).

Participants performed a highly-demanding, short-term memory task. On some trials they relied on internal memory alone (which typically resulted in poor accuracy); on other trials they could set external reminders (which dramatically improved accuracy). The key manipulation was the number of points associated with the two strategies. Correct responses always earned 10 points if participants used internal memory, but a lower number of points between 2-9 if they used external reminders. The latter number of points varied from trial to trial, and participants were required each time to decide which strategy they preferred (e.g., 10 points for each correct response with internal memory, or 6 points for each correct response with external reminders). The ‘optimal indifference point’ was that point value at which an unbiased individual would be indifferent between the two strategies, based on their objective accuracy in the two conditions. The ‘actual indifference point’ was the point at which they were actually indifferent, based on all of their decisions. By comparing these two values, we obtained a ‘reminder bias’: the extent to which an individual had a pro- or anti-reminder bias relative to their individually-calculated optimal strategy. Note that this is different from the absolute rate of reminder usage, because the same absolute rate might reflect inadequate use of reminders in a person with poor memory and excessive reminder usage in a person with good memory ability. Along with the reminder bias, we also calculated a metacognitive bias, which represents participants’ over- or underconfidence in memory ability, relative to objective performance. Our study controlled for age, gender, educational attainment, as well as cognitive ability (ICAR5; Kirkegaard & Bjerrekær, 2016), and working memory.

Previewing our results, in line with previous evidence, we found that confidence varied positively with the CIT factor and negatively with the AD factor. However, contrary to our initial expectations, more compulsive individuals offloaded less rather than more, and there was no evidence for disruption in the link between metacognition and offloading. Instead, we discovered an incomplete mediation effect: while a significant proportion of the reduced reminder setting could be attributed to overconfidence, not all the variance was accounted for by this variable. Even after controlling for it, compulsivity still predicted reduced reminder setting. This constitutes a combination of the Metacognitive Monitoring Mechanisms and the Direct Mechanism.

Results

Here we present the results of a preregistered online study on the relationship between reminder setting, metacognition and transdiagnostic compulsivity. We excluded 69 out of a total of 669 participants based on our six preregistered criteria described in the Methods and Materials section, leaving us with a final sample of 600 participants. All participants completed a previously validated reminder setting task in combination with 49 items from six mental health questionnaires. Three-hundred-and-seventy-five participants identified as male, 218 as female and 7 as other. Participants were on average 32.9 years old (min = 18; max = 76). Figure 1 shows the included (black) and excluded (red) data, with higher average performance for included participants when reminders were used (96.1%) compared to when people had to do the task unaided (59.2%).

We calculated six key measures for each participant:

The first relevant measure is the optimal indifference point (OIP). The OIP describes the reward value (2-9 points) at which an unbiased, reward-maximising participant should be indifferent between the two strategies: using reminders or relying on their own memory. The OIP is calculated from their accuracy with and without reminders. Imagine a participant who achieves 60% accuracy when using their own memory or 100% accuracy when using reminders. In this case the OIP would be 6, because scoring 6 points per item with reminders (100% accuracy) would earn the same number of points as scoring 10 points with internal memory (60% accuracy). For any reward above 6, it would be optimal to choose external reminders; for any reward below 6 it would be optimal to choose internal memory.

In contrast, the second relevant measure is the actual indifference point (AIP), which is the number of points at which participants showed indifference between the two strategies. This measure is calculated by fitting a psychometric function to participants’ choices at different levels of reward for targets when reminders were used. Please note that all choices were used to calculate the AIP, as participants only found out whether or not they would use a reminder after the decision was made.

Together, these variables can be used to calculate the third measure, the reminder bias, which is the difference between the OIP and the AIP and therefore reflects participants’ tendency to over- or under-use reminders, relative to the optimal strategy. Note that the optimal strategy is calculated individually for each participant and will depend on their own level of performance when using internal memory and external reminders.

Fourth, we calculated a metacognitive bias, reflecting participants’ over- or underconfidence. This is calculated by subtracting objective accuracy (percentage of targets remembered when using internal memory) from the percentage that they predicted that they would be able to remember.

Fifth and sixth, based on the questionnaire ratings we calculated how much someone scored on the transdiagnostic CIT and AD factors. Our analyses focus on the relationship between these key measures.

Added in the revised version: Average accuracy as a function of whether a reminder was used.

Earlier version: Average accuracy as a function of whether a reminder was used. “No Reminder”: forced internal condition; “Reminder”: forced external condition. Each pair of dots linked by a line indicates one participant. The red data points are excluded participants. The box plots indicate the median surrounded by the interquartile range (25th and 75th percentile). The whiskers show the minimum and maximum. The preregistered exclusion criteria for the accuracies with or without reminder are indicated as horizontal dotted lines (10% and 70% respectively).

Revised version: “No Reminder”: forced internal condition; “Reminder”: forced external condition. Each pair of dots linked by a line indicates one participant. The red data points are excluded participants. The box plots indicate the median surrounded by the interquartile range (25th and 75th percentile). The whiskers show the minimum and maximum. The preregistered exclusion criteria for the accuracies with or without reminder are indicated as horizontal dotted lines (10% and 70% respectively).

Replication and sanity checks

Earlier version: In the following section, we aim, where the design allows it, to replicate four previous effects for this task. First, with Hypothesis 1, we predicted that the reminder bias and metacognitive bias are negatively correlated replicating previous findings (as reviewed in Gilbert et al., 2023). This effect tests the above-mentioned link between metacognition and cognitive offloading: the less confident someone feels, the more they use reminders. There was indeed a significant negative correlation, r = −0.2, p <.001 (Figure 2). Second, in replication of previous findings (e.g., Gilbert et al., 2020; Sachdeva & Gilbert, 2020; Kirk, Robinson & Gilbert, 2021; Engeler & Gilbert, 2020), Hypothesis 2 expressed our expectation to find an excessive use of reminders reflected in significantly higher optimal indifference points (OIPs) compared to actual indifference points (AIPs). In other words, we expected the reminder bias to be greater than zero, which was indeed the case, m = 0.52, t = 5.1, p <.001, d = 0.21. Third, with Hypothesis 3, we expected to replicate that participants would be underconfident in their own memory (e.g., Engeler & Gilbert, 2020), expressed in an average, negative metacognitive bias. Our data supported this hypothesis, m = −3.64, t = −3.1, p = .001, d = −0.13. Fourth, Hypothesis 4 predicts that as in previous studies, we would find evidence for compensatory reminder use. Keeping in mind that the OIP reflects the cut-off at which participants should be indifferent between offloading and not offloading and the AIP the cut-off they actually displayed, then looking at these two measures together should show that participants with poorer memory and greater benefit from reminders (lower OIP) tend to use them more (lower AIP). Indeed, the optimal and actual indifference points were positively correlated suggesting participants who benefited most from reminders were more likely to use them, r = 0.36, p <.001. Taken together, we found that participants showed the usual hallmarks of this offloading task, using their confidence to strategically decide when to offload, general tendencies for setting reminders and for underconfidence, and compensatory reminder use.

Revised version: In the following section, we aim, where the design allows it, to replicate four previous effects for this task. First, with Hypothesis 1, we predicted that the reminder bias and metacognitive bias are negatively correlated replicating previous findings (as reviewed in Gilbert et al., 2023). This effect tests the above-mentioned link between metacognition and cognitive offloading: the less confident someone feels, the more they use reminders. There was indeed a significant negative correlation, r = −0.2, p <.001 (Figure 2). Second, in replication of previous findings (e.g., Gilbert et al., 2020; Sachdeva & Gilbert, 2020; Kirk, Robinson & Gilbert, 2021; Engeler & Gilbert, 2020), Hypothesis 2 expressed our expectation to find an excessive use of reminders reflected in significantly higher optimal indifference points (OIPs) compared to actual indifference points (AIPs). In other words, we expected the reminder bias to be greater than zero, which was

Added in the revised version: indeed the case, m = 0.52, t = 5.1, p <.001, d = 0.21. Third, with Hypothesis 3, we expected to replicate that participants would be underconfident in their own memory (e.g., Engeler & Gilbert, 2020), expressed in an average, negative metacognitive bias. Our data supported this hypothesis, m = −3.64, t = −3.1, p = .001, d = −0.13. Fourth, Hypothesis 4 predicts that as in previous studies, we would find evidence for compensatory reminder use. Keeping in mind that the OIP reflects the cut-off at which participants should be indifferent between offloading and not offloading and the AIP the cut-off they actually displayed, then looking at these two measures together should show that participants with poorer memory and greater benefit from reminders (lower OIP) tend to use them more (lower AIP). Indeed, the optimal and actual indifference points were positively correlated suggesting participants who benefited most from reminders were more likely to use them, r = 0.36, p <.001. Taken together, we found that participants showed the usual hallmarks of this offloading task, using their confidence to strategically decide when to offload, general tendencies for setting reminders and for underconfidence, and compensatory reminder use.

Earlier version: People’s tendency to set reminders above or below the optimal offloading strategy (reminder bias) plotted against people’s tendency towards over- or underconfidence (metacognitive bias). The solid line indicates the fitted relationship between both variables. The dashed line represents the 95% confidence interval around it. Each circle represents a single participant.

Revised version: People’s tendency to set reminders above or below the optimal offloading strategy (reminder bias) plotted against people’s tendency towards over- or underconfidence (metacognitive bias).

Added in the revised version: The solid line indicates the fitted relationship between both variables. The dashed line represents the 95% confidence interval around it. Each circle represents a single participant.

Testing our key hypotheses

Elevated confidence in CIT and reduced confidence in AD

We predicted that the metacognitive bias would correlate negatively with AD (Hypothesis 8a; more anxious-depressed individuals tend to be underconfident). For CIT, we preregistered a non-directional, significant link with metacognitive bias (Hypothesis H6a). We found support for both hypotheses, both for AD, β = −0.22, SE = 0.04, t = −5.00, p < 0.001, as well as CIT, β = 0.15, SE = 0.05, t = 3.30, p = 0.001, controlling for age, gender, and educational attainment (Figure 3; see also Table S1). Note that for CIT this effect was positive, more compulsive individuals tend to be overconfident.

Earlier version: Standardised regression weights for the ‘anxious-depression’ factor (AD) and the ‘compulsive behaviour and intrusive thought’ factor (CIT) predicting metacognitive bias. Error bars indicate 95% confidence intervals. Asterisks indicate significance: “***”: < 0.001; “**”: < 0.01; “*”: < 0.05.

Revised version: Standardised regression weights for the ‘anxious-depression’ factor (AD) and the ‘compulsive behaviour and intrusive thought’ factor (CIT) predicting metacognitive bias. Error bars indicate 95% confidence intervals.

Added in the revised version: Asterisks indicate significance: “***”: < 0.001; “**”: < 0.01; “*”: < 0.05.

We furthermore preregistered to also test this for raw confidence (percentage of circles participants predicted they will remember, rather than the accuracy-corrected metacognitive bias score; Hypotheses H8b and H6b). Indeed, the same patterns were found for both AD, β = − 0.29, SE = 0.04, t = −6.44, p < 0.001, and CIT, β = 0.16, SE = 0.05, t = 3.50, p < 0.001 (see Table S2). Including scores from the cognitive ability test as an additional covariate (Hypotheses H8c and H6c respectively) furthermore did not change the results, AD, β = −0.20, SE = 0.04, t = −4.47, p < 0.001; CIT, β = 0.12, SE = 0.04, t = 2.74, p = 0.006 (see Table S3). Taken together, these results suggest that concordant with our hypotheses, compulsivity was linked to inflated confidence and anxiety to deflated confidence.

Contrary to Expectations, Compulsivity Reduced Pro-offloading Bias

We expected to find a positive link between CIT factor scores and reminder bias. In other words, we predicted that more compulsive individuals would show a greater pro-offloading bias, relative to the optimal strategy (Hypothesis H5a). However, our results showed the exact opposite effect with a significantly reduced reminder bias in compulsive individuals, β = −0.16, SE = 0.04, t = −3.60, p < 0.001, controlling for age, gender, and educational attainment (Figure 4; see also Table S4). This trend persisted when, instead, we predicted the absolute number of reminders chosen by the participant (Hypothesis H5b), β = −0.14, SE = 0.05, t = −3.08, p = 0.002 (see Table S5), as well as when predicting the AIP (Hypothesis H5c), β = 0.15, SE = 0.05, t = 3.32, p < 0.001 (see Table S6).

Earlier version: Standardised regression weights for the ‘anxious-depression’ factor (AD) and the ‘compulsive behaviour and intrusive thought’ factor (CIT) predicting reminder bias. Error bars indicate 95% confidence intervals. Asterisks indicate significance: “***”: < 0.001; “**”: < 0.01; “*”: < 0.05.

Revised version: Standardised regression weights for the ‘anxious-depression’ factor (AD) and the ‘compulsive behaviour and intrusive thought’ factor (CIT) predicting reminder bias.

Added in the revised version: Error bars indicate 95% confidence intervals. Asterisks indicate significance: “***”: < 0.001; “**”: < 0.01; “*”: < 0.05.

Previous studies have found reduced working memory in OCD (Harkin & Kessler, 2011), which could potentially lead to increased reminder use in compulsivity. However, the reduced reminder bias persisted if d’ from the 2-back task was included as an additional covariate (Hypothesis H5d), β = −0.15, SE = 0.05, t = −3.27, p = 0.001 (see Table S7). Finally, we predicted that our results would persist independent of whether or not the scores from the cognitive ability test were included as an additional covariate (Hypothesis H5e), which was indeed the case, β = −0.16, SE = 0.05, t = −3.55, p < 0.001 (see Table S8). It should be noted that all our regression models included both CIT and AD as predictors to separate out the potentially competing influences of these predictors, as well as age, gender, and educational attainment as demographic covariates.

We furthermore preregistered to conduct the same tests for the AD factor but without any directional hypotheses. AD was not significantly linked to any changes in reminder bias, β = 0.06, SE = 0.04, t = 1.43, p = 0.15 (see Table S4), absolute number of reminders, β = 0.06, SE = 0.05, t = 1.32, p = 0.19 (see Table S5), or AIP, β = −0.08, SE = 0.04, t = −1.74, p = 0.08, controlling for age, gender, and educational attainment (see Table S6). This null effect did not change when working memory, β = 0.05, SE = 0.04, t = 1.21, p = 0.23 (see Table S7), or scores from the cognitive ability test were included as additional covariates, β = 0.06, SE = 0.04, t = 1.40, p = 0.16 (see Table S8).

Taken together, these results suggest that compulsive individuals are less biased towards offloading, in contrast to our hypothesised direction of the effect, but consistent with the observation of increased confidence in their ability on this task.

No Evidence for Impaired Confidence-Offloading Link

We predicted to find support for the Metacognitive Control Mechanism, meaning that CIT would act as a moderator on the link between confidence and offloading (Hypothesis H7a). In other words, we expected to find that the correlation between the metacognitive and the reminder bias to be weakened in highly compulsive individuals. However, the interaction between metacognitive bias and compulsivity in a model predicting the reminder bias was not significant, β = −0.01, SE = 0.04, t = −0.18, p = 0.86, controlling for age, gender, and educational attainment (see Table S9). This means that in our task, confidence and offloading were linked just as much as in their low compulsive counterparts. These results remained the same even if working memory performance (d’ from the 2-back task) was included as an additional covariate (Hypothesis H7b), β = −0.01, SE = 0.04, t = −0.26, p = 0.80 (see Table S10), or if scores from the cognitive ability test were included as an additional covariate (Hypothesis H7c), β = −0.01, SE = 0.04, t = −0.18, p = 0.86 (see Table S10).

Contrary to our initial hypotheses, we found that increased CIT was associated with decreased rather than increased bias towards offloading. Seeing as CIT was also associated with increased confidence, and high confidence predicts low bias towards offloading, we tested whether the relationship between CIT and offloading was mediated via confidence: whilst parts of the reduction of reminders could be traced back to overconfidence, β = −0.18, SE = 0.04, t = −4.51, p < 0.001, there was still a significant proportion of variance that was linked to compulsivity independently of this, β = −0.14, SE = 0.05, t = −3.03, p = 0.003. Figure 5 summarises this incomplete mediation effect. In previous sections we have already reported the total effect of CIT on the reminder bias, β = −0.16, SE = 0.05, t = −3.60, p < 0.001. Equally, we have already reported the effect of CIT on the mediator (the metacognitive bias), β = 0.12, SE = 0.05, t = 2.76, p = 0.006.

Earlier version: Diagram of the mediation analysis testing for the influence of the ‘compulsive behaviour and intrusive thought’ factor (CIT) on reminder bias, both directly and indirectly through the metacognitive bias. Standardised regression coefficients are given for each path. The value in parentheses indicates the influence of CIT on reminder bias controlling for the influence of the metacognitive bias. Asterisks indicate significance: “***”: < 0.001; “**”: < 0.01; “*”: < 0.05.

Revised version: Diagram of the mediation analysis testing for the influence of the ‘compulsive behaviour and intrusive thought’ factor (CIT) on reminder bias, both directly and indirectly through the metacognitive bias.

Added in the revised version: Standardised regression coefficients are given for each path. The value in parentheses indicates the influence of CIT on reminder bias controlling for the influence of the metacognitive bias. Asterisks indicate significance: “***”: < 0.001; “**”: < 0.01; “*”: < 0.05.

We validated this outcome through an exploratory causal mediation analysis. The indirect influence of the CIT factor on reminder bias going through the metacognitive bias was calculated to be (0.12) * (-0.18) = −0.02. To determine the significance of this influence, we implemented bootstrapping procedures. We computed unstandardized indirect effects for each of the 1,000 bootstrapped samples, followed by the calculation of the 95% confidence interval, identifying the indirect effect at the 2.5th and 97.5th percentiles. The bootstrapped unstandardized indirect effect (the average causal mediation effect) computed to be −0.02, with the 95% confidence interval ranging from −0.05 to −0.01. This indicated that the effect was statistically significant at p < 0.001.

Finally, we preregistered to run the same analysis for the AD factor without hypothesising about any specific direction for any potential effects. However, we did not find evidence for a moderation effect (an interaction between AD scores and metacognitive bias when predicting the reminder bias), β = −0.04, SE = 0.04, t = −1.13, p = 0.26, controlling for age, gender, and educational attainment (see Table S12).

In summary, whilst we found no support for the Metacognitive Control Mechanism (as would be reflected in a disrupted link between confidence and offloading), we did find support for both the Metacognitive Monitoring Mechanism (reduced pro-reminder bias as a downstream consequence of overconfidence) and the Direct Mechanism (independent contribution of CIT on offloading). The Supplement furthermore lists several additional analyses, both planned and exploratory.

Discussion

In the current study, we explored the behavioural and cognitive correlates of two transdiagnostic traits: ‘compulsive behaviour and intrusive thought’ (CIT) and ‘anxious-depression’ (AD). We focused on changes in cognitive offloading and metacognition related to transdiagnostic compulsivity. Our results replicated that more compulsive individuals were relatively overconfident, while those who were more anxious-depressed were relatively underconfident. Contrary to expectations, we observed a decreased bias towards reminders among more compulsive participants. This reduction in bias was only partially accounted for by their relative overconfidence. This partial mediation can be interpreted through both a Metacognitive Monitoring Mechanism (differences in the formation of the confidence signal rather than its behavioural application) and a Direct Mechanism (no metacognitive involvement). We found no support for a Metacognitive Control Mechanism, which would centre on how confidence is used to adapt behaviour (Nelson & Narens, 1990; Boldt & Gilbert, 2022).

Perfectionism and the need to control as potential explanations

Contrary to our hypothesis, our study revealed an inverse relationship between transdiagnostic compulsivity and offloading: the reminder bias was reduced in more compulsive individuals. One possible interpretation is perfectionism: Some compulsive individuals may avoid using reminders altogether due to rigid, perfectionistic beliefs about needing to remember everything without relying on external aids, and using reminders could trigger their anxiety or feed into their obsessions about being forgetful or unreliable. This interpretation aligns with findings suggesting that perfectionism serves as a transdiagnostic maintaining and risk factor for various mental health conditions, including compulsive disorders like eating disorders and OCD (Egan, Wade, & Shafran, 2011).

No effect of anxiety on offloading

Interestingly, we found no significant influence of the anxious-depressed (AD) transdiagnostic phenotype on offloading. This aligns with a recent study by Kirk, Robinson, & Gilbert (2021), which also found no effect of anxiety on offloading. However, their study, which used the ‘trait’ component of the STAI to measure anxiety (Spielberger, Gorsuch, Lushene, Vagg, & Jacobs, 1983), found no relative underconfidence among anxious participants either. Our transdiagnostic approach likely revealed this confidence effect by separating the counteracting influences of AD and CIT factors. This distinction underscores the value of a transdiagnostic approach.

Our findings align with those reported in a recent study by Mohr, Ince, and Benwell (2024). The authors observed that while high-AD participants were underconfident in a perceptual task, this underconfidence did not lead to increased information-seeking behaviour. Future research should explore whether this is due to their pessimism regarding the effectiveness of confidence-modulated strategies (i.e., setting reminders or seeking information) or whether it stems from apathy. Another possibility is that the relevant downstream effects of anxiety were not measured in our study and instead may lie in reminder-checking behaviours.

No evidence for an impaired confidence-action link in compulsivity

Contrary to Seow and Gillan (2020) and Vaghi et al. (2017), our study did not find the impaired confidence-action link (Metacognitive Control Mechanism) reported for transdiagnostic compulsivity and OCD patients on this task. This may be because of differences between tasks – prior work used a reinforcement learning task with a clear learning element from trial to trial. Alternatively, it is possible our study was underpowered, as our sample size was designed to detect overconfidence in compulsivity, not the more nuanced but still psychometrically robust confidence-action link (Loosen, Seow & Hauser, 2023), which would have required a far larger sample size. Recent studies also failed to find decreased action-confidence coupling with relatively small groups of OCD patients and controls (Hoven et al., 2023b; Marzuki et al., 2022). Indeed, both our paradigm and the earlier predictive-inference task tested for an interaction effect, which is more challenging to power adequately. Future research should consider using more direct measures that ideally aim to manipulate confidence directly.

Implications

Participants in our current study were recruited from the general population through Prolific, meaning that the variance likely represents primarily subclinical sources. Consequently, caution should be exercised when extrapolating these results to clinical populations. For example, a recent study indicated that metacognitive impairments in OCD originate from different mechanisms than those observed in transdiagnostic compulsivity (Hoven et al., 2023c). Given its metacognitive impairments and the prevalent symptom of checking, OCD still remains a particularly relevant patient group for studying reminder setting and future studies need to explore this area further. Due to their underconfidence, OCD patients might engage in more frequent reminder setting. This behaviour could serve as a compensatory mechanism, especially since OCD patients often face challenges with working memory (Harkin & Kessler, 2011) and prospective memory (Harris, Vaccaro, Jones, & Boots, 2010; Racsmány, Demeter, Csigó, Harsányi, & Németh, 2011). However, it could also worsen their checking symptoms as more reminders means more opportunities to check.

On the other hand, it is possible that the observed underconfidence in OCD populations may actually reflect the impact of an uncontrolled anxiety factor, effectively neutralizing the influence of compulsivity on confidence. This confounding issue could explain the inconsistent findings regarding confidence bias in both compulsivity and OCD. If this was the case, then future research should investigate which influences on confidence – the reductions caused by the AD factor or the increases caused by the CIT factor – are the driving force behind any changes in reminder setting in OCD.

A pivotal question remains: will the overall reduction in reminder setting, referred to as a “Direct Effect” in this study, also be observed in OCD patients and other compulsive disorders? Such findings could support the hypothesis that an inherent aspect of compulsivity leads to the decreased use of external aids, potentially due to perfectionism or a need for control.

Limitations

Our results are based on a well-validated paradigm which our lab has previously used on other, published studies (as reviewed in Gilbert et al., 2023). However, reliance on a single behavioural task also means that our results might not generalise onto cognitive offloading more broadly or even reminder setting in other contexts. As a first step, future work should aim to replicate our findings in the context of other experimental designs.

Another limitation is that in the present study, we focused solely on measuring two transdiagnostic factors: CIT and AD. We omitted the third factor, ‘social withdrawal’ (SW). By doing so, we were able to reduce the number of items from six clinical questionnaires to 49 (Wise & Dolan, 2020), thereby shortening the required time for completion—an essential consideration for online research (Sauter, Draschkow & Mack, 2020). Nevertheless, this focused approach could introduce variability in capturing these transdiagnostic phenotypes. A recent preprint from Hopkins and colleagues (2022) supports this approach. They used machine learning to select 71 items capable of reliably measuring all three factors, suggesting that future transdiagnostic studies might similarly adopt more concise item sets.

Conclusion

With the present study, we investigated the downstream cognitive and behavioral effects of two transdiagnostic traits, CIT and AD. In particular, we were interested in the effect these factors have on metacognition and cognitive offloading, operationalised as prospective confidence and reminder setting respectively. We replicated the finding that more compulsive individuals tend to be relatively overconfident, whereas anxious-depressed individuals tend to be relatively underconfident. Contrary to our hypotheses, however, we found that compulsivity was linked reduced offloading and that this effect was only in part explained by overconfidence.

Fulfilling delayed intentions (i.e., prospective memory) is a vital process for daily living and behavioral independence. However, this process is also highly fallible (e.g., Crawford, Smith, Maylor, Della Sala & Logie, 2003). External memory aids are highly effective and commonplace tools that compensate for these memory failures (e.g., Jones, Benge & Scullin, 2021; Scullin et al., 2022). Our findings suggest that compulsive individuals are at particular risk of inadequate external memory support and would potentially benefit from interventions that target cognitive offloading strategies.

Methods

Task and Procedure

For the present, preregistered study, we used a novel variant of an online cognitive-offloading task (“optimal reminders task”; cf. Gilbert et al., 2020). This task allowed us to measure how people set reminders in relation to their confidence. All procedures, hypotheses, and planned analyses were preregistered at https://osf.io/kztf8 prior to the commencement of data collection.

On every trial, participants were instructed to move several numbered, yellow circles to the bottom of a square in consecutive order (see Figure 6A). Whenever a circle was removed, a new one appeared up to a total of 15 circles. The source of difficulty of this task stems from the ‘special’ circles, which constitute the delayed intentions people have to fulfil. These circles flashed in a colour (blue, orange, or magenta) when they first appear on screen before fading to yellow. Participants’ task was to drag these circles to their colour-corresponding side once the time has come to remove the respective special circle (top, left or right). There were 6 special circles per trial. On some trials, participants had to rely on their own memory to complete the task and remember the target locations of the special circles. On other trials, they set spatial reminders indicating the locations to which the special circles must be moved to. More specifically, they were taught to move the special circle next to the border through which it would have to be moved out of the square later.

A) Example sequence of events within a single trial. Trajectories of movement made by a fictive participant are shown as black arrows. The blue coloured circle corresponds to the left boundary of the square and indicates that this circle must be moved to this side rather than the bottom. B) Example of an offloading decision which participants were required to make before each trial. C) After each decision, they were informed whether or not they would perform the upcoming trial with reminders. The cell’s shading indicates the participant’s original choice. D) Confidence was rated once before the introduction of the offloading strategy on a scale ranging from 0 to 100%. E) Sequence of events within the task. All aspects of the task were performed online in the web browser.

Every trial began with a decision: participants could choose to do the task without reminders and earn 10 points for every special circle they remembered to move to the correct border, or they could choose to use reminders but earn less for each special circle (Figure 6B). Critically, this lesser amount was varied between 2 and 9 points, allowing us to calculate participants’ indifference point when trading off the benefit of reminders with their reduced reward. This actual indifference point (AIP) could then be contrasted against their optimal indifference point (OIP), calculated from participants’ accuracy with or without reminders, see below for further details. Since our task included only 4 trials each with or without reminders, we counterbalanced the assignment of odd or even target values to these conditions.

Together, there were three key conditions in our task, presented intermixed throughout the experiment: the Forced Internal condition (FI; 4 trials) in which participants had to remember the circles unaided, the Forced External condition (FE; 4 trials) in which they had to use reminders, and the Choice Only (CO; 8 trials) condition in which they were free to choose whichever strategy they preferred but the trial ended after only six circles and without any special circles. To give participants the impression of maximum agency over the task, we only told them that their choice would be overwritten whenever there was a mismatch with the pseudorandomly assigned condition (25.3% of all trials; SD = 5.8; see Figure 6C). This way, participants were unable to tell which condition they were currently in and whether it would be a partial trial. Participants used reminders on average on 49.9% of trials (SD = 16.5).

Participants were asked to rate their confidence once during the experiment, being asked to indicate the “percentage of the special circles [they] can correctly drag to the instructed side of the square” (Figure 6D). Importantly, this confidence judgement was given after the first practice trials and before the offloading strategy was introduced to ensure participants answered this question with regard to their own perceived memory capabilities. Average confidence was 55.6% (proportion of trials on which participants predicted to remember to move the special circles; SD = 24.2).

In addition to the reminder task, we included items from six individual differences questionnaires shortened to include only the items required to reliably measure the CIT and AD factors (Wise & Dolan, 2020). These questionnaires were presented in random order: 4 items from the Apathy Evaluation Scale (AES; Marin, Biedrzycki, & Firinciogullari, 1991), 8 items from the Zung Depression Scale (SDS; Zung, 1965), 4 items from the Eating Attitudes Test (EAT-26; Garner, Olmsted, Bohr, & Garfinkel, 1982), 12 items from the Barratt Impulsiveness Scale (BIS-11; Patton, Stanford, & Barratt, 1995), 11 items from the Obsessive Compulsive Inventory – Revised (OCIR; Foa, Huppert, Leiberg, Hajcak, Langner, et al., 2002), and 11 items from the ‘trait’ part of the State-Trait Anxiety Inventory (STAI; Spielberger, Gorsuch, Lushene, Vagg, & Jacobs, 1983). A list of all included items can be found in the Supplement.

We furthermore included a catch item in the BIS-11 (“I competed in the 1917 Summer Olympics Games.”) to ensure participants were paying attention to the task as well as three covariates aimed at measuring cognitive ability (a 5-item version of the International Cognitive Ability Resource; ICAR5; Kirkegaard & Bjerrekær, 2016; Condon & Revelle, 2014); educational attainment mapped onto a 1 to 9 scale and based on the ISCED 2011 categories (see Supplement), and working memory (100 consecutive letters of the 2-back task; e.g. Kirchner, 1985). The logic behind including the latter covariate was that whilst our key dependent variables already corrected for working memory (more specifically: unaided prospective memory performance) this could tap into additional working memory components not measured already and potentially impacted in compulsivity based on the finding that they have often been found to be impaired in OCD (Harkin & Kessler, 2011). Together, these elements resulted in a total duration of approximately 35 mins. The sequence of events within the task is shown in Figure 6E.

Participants

Participants were invited on prolific.co to participate for £3.90. Based on points won during the main task, the upper 50% of participants were furthermore rewarded with a bonus payment of £1. We restricted our search to the Prolific standard sample, allowed participants from all countries, with a minimum of 18 years. All participants had to be fluent in English and were required to have an approval rate of over 90% based on Prolific’s criteria. Moreover, we required participants to not have participated in one of the four pilots prior to this study.

We calculated our sample size based on the link between confidence and transdiagnostic compulsivity as reported in two recent studies (Rouault et al., 2018; Seow & Gillan 2020). To be able to detect a link between these variables of β = 0.23, p < 0.001, as in Rouault et al. (2018), we required N = 288 participants (two-sided testing, power = 0.8, CL = 0.95). To be able to detect a link of β = 6.74, p < 0.001, as in Seow and Gillan (2020), we required N = 291 participants (two-sided testing, power = 0.8, CL = 0.95). In both cases, the power calculation was based on a partial regression approach, excluding the effect in question from the model and comparing the explained variance compared to the full model. Since we are furthermore aiming to test a moderation effect of compulsivity on the link between the metacognitive bias and the reminder bias, we decided to collect a larger sample of N = 600 after exclusions.

Earlier version: We preregistered six exclusion criteria, based on which we excluded and replaced 69 participants: Nine participants were excluded due to a higher hit rate on forced internal than forced external trials, 22 participants were excluded due to less that 70% accuracy on FE trials, and 3 participants due to less than 10% accuracy on FI trials. We furthermore preregistered to exclude participants with a negative correlation between value and reminder choice (1 = reminder, β = no reminder), as this would indicate participants did not understand the instructions: in order to maximise points in our task, participants should preferentially choose reminders when this strategy brings a higher number of points. Based on this, we excluded 40 participants. No participants were excluded based on scoring lower or higher 3 times the median absolute deviation (MAD), calculated separately based on both the reminder bias and the metacognitive bias. Finally, we excluded 9 participants because they failed to answer with “Do not agree at all” to the catch item. Figure 1 visualizes the exclusions shown in red. In total, we excluded 10.3% of all participants. There were an additional 26 participants excluded for technical reasons, raising the exclusion rate to 13.7%.

Revised version: We preregistered six exclusion criteria, based on which we excluded and replaced 69 participants: Nine participants were excluded due to a higher hit rate on forced internal than forced external trials, 22 participants were excluded due to less that 70% accuracy on FE trials, and 3 participants due to less than 10% accuracy on FI trials. We furthermore preregistered to exclude participants with a negative correlation between value and reminder choice (1 = reminder, 0 = no reminder), as this would indicate participants did not understand the instructions: in order to maximise points in our task, participants should preferentially choose reminders when this strategy brings a higher number of points. Based on this, we excluded 40 participants. No participants were excluded based on scoring lower or higher 3 times the median absolute deviation (MAD), calculated separately based on both the reminder bias and the metacognitive bias. Finally, we excluded 9 participants because they failed to answer with “Do not agree at all” to the catch item. Figure 1 visualizes the exclusions shown in red. In total, we excluded 10.3% of all participants. There were an additional 26 participants excluded for technical reasons, raising the exclusion rate to 13.7%.

Key dependent variables

Our task allowed us to calculate several dependant variables relevant in the context of our study question. The first is the OIP, the optimal indifference point. The OIP describes the number of points at which an unbiased, reward-maximising participant is indifferent between the two strategies (reminders or no reminders) and is calculated as: OIP=(10*ACCFI)/ACCFE

Earlier version: Where ACCFE is the accuracy measured during trials in which participants had to solve the task using reminders (FE condition) and ACCFI is the accuracy measured during trials in which participants had to solve the task without reminders (FI condition). In contrast, the AIP is the actual indifference point, which is the point cut-off at which participants actually were indifferent and is operationalised as the threshold parameter α from fitting a psychometric function to the choice data (target values predicting the decision whether or not to use reminders). Fitting was done using the quickpsy package in R.

Revised version: Where ACCFE is the accuracy measured during trials in which participants had to solve the task using reminders (FE condition) and ACCFI is the accuracy measured during trials in which participants had to solve the task without reminders (FI condition). In contrast, the AIP is the actual indifference point, which is the point cut-off at which participants actually were indifferent and is operationalised as the threshold parameter from fitting a psychometric function to the choice data (target values predicting the decision whether or not to use reminders). Fitting was done using the quickpsy package in R and more detail is given in the Supplement.

Setting the OIP and the AIP in relation, we can calculate the reminder bias, reflecting participants’ tendency to use reminders corrected for their actual performance and calculated as the difference between both indifference points: biasrem=OIP−AIP

Positive values reflecting that people set more reminders relative to the optimal strategy. The fourth measure is the metacognitive bias, reflecting participants’ over- or underconfidence relative to their performance and was calculated as: biasmeta=confidence−ACCFI

Negative values can be interpreted as underconfidence.

Crucially, our study relies on the key assumption that the metacognitive bias can predict the reminder bias, but ACCFI contributes to both biases. To avoid circularity, we therefore split the accuracy data to avoid potentially inflating the correlation. More specifically, we included only the even trials to calculate the ACCFI for the OIP, whereas we included only the odd trials to calculate the ACCFI for the metacognitive bias. All available trials from the FE condition were used to calculate the OIP.

It should be noted that we had incorrectly stated in the preregistration that accuracy from forced external trials would contribute to the calculation of the metacognitive bias. However, the metacognitive bias is a judgement given about the unaided memory performance, in fact confidence is measured before participants were even introduced to the offloading strategy (see above). We therefore used only the internal trials in calculating the metacognitive bias.

Finally, the transdiagnostic scores for the ‘compulsive behaviour and intrusive thought’ (CIT) factor and the ‘anxious-depression’ (AD) factor were calculated from participants’ ratings to the individual differences questionnaires by multiplying them with the item weights from Wise and Dolan (2020) prior to summing them. The items composing the CIT and AD scores respectively were non-overlapping with 24 items forming the AD score and 25 items forming the CIT score.

Preregistered Hypotheses and Statistical Analyses

We preregistered 8 hypotheses (see Table 1), half of which were sanity checks (H1-H4) aimed to establish whether our task would generally lead to the same patterns as previous studies using a similar task (as reviewed in Gilbert et al., 2023). H1 was a replication of the central finding of the link between confidence and offloading. More specifically, we entered the unconfounded metacognitive bias and reminder bias into a Pearson correlation analysis. We expected to find a negative relationship between the two measures, which we planned to test for significance using a one-sided test. We furthermore expected to find that people would use more reminders than optimal. This pro-reminder bias would be reflected in a positive reminder bias (H2). We planned to test this using a one-sided paired t-test. Relatedly, we expected to find people to be generally underconfident (i.e., expecting to remember fewer special circles than they actually did when doing the task without reminders). Such underconfidence would be reflected in a negative metacognitive bias (H3), which we again planned to test this using a one-sided paired t-test. Furthermore, we expected that those who required more reminders would also be the ones to use them more, as reflected in a positive correlation between the AIP and OIP, again as a one-sided test (H4). We decided to use Spearman’s rho due to the data most likely being distributed around the extremes of the scale. For H2-H4 (as well as H5, H6, and H8, see below), we used the biases and indifference points calculated from all available trials as there was no circularity issue.

Hypotheses H5-H8 were the key hypotheses of our study. Here, we address them out of order in the interest of an improved logical flow. Hypothesis H6 predicted that more compulsive individuals would show a confidence bias, reflected in a reliable predictor of the CITs scores on the metacognitive bias from the following regression model: biasmeta~CIT+AD+age+gender+education+ε

Though we did not preregister a direction for this effect, in the light of recent findings it has now become clear that compulsivity would most likely be linked to overconfidence (Rouault et al., 2018; Seow & Gillan, 2020; Benwell et al., 2022; Fox et al., 2023a; Fox et al., 2023b; Hoven et al., 2023a). The same model was used to test hypothesis H8, predicting that more anxious-depressed individuals tend to be underconfident. This would be reflected in AD scores being negatively linked to the metacognitive bias. The model above represents the main models designed to test hypotheses H6a and H8a. We furthermore also tested these hypotheses but predicted raw confidence (percentage of circles participants predicted they would remember; H6b and H8b respectively), as well as extending the main model with the scores from the cognitive ability test (ICAR5) as an additional covariate (H6c and H8c respectively). For this as well as all following regression models, we z-transformed all non-binary variables prior to fitting the models.

With H5, we predicted that more compulsive individuals would show a bias towards more offloading, reflected in a positive regression coefficient when using the CIT score as a predictor of the reminder bias. This hypothesis was not a replication, consequently we decided to carry out the test two-sided. Throughout this section, whenever not explicit specified, we plan to carry out a test two-sided. Due to the diametrically opposing effects of CIT and AD, both transdiagnostic scores need to be entered into the model, alongside our demographic covariates age, gender, and educational attainment: biasrem~CIT+AD+age+gender+education+ε

We fitted several different versions of this model: the main model predicted the reminder bias (H5a), but we also fit one with the absolute number of reminders chosen (H5b), or the AIP (H5c). To understand whether any differences in offloading behaviour could stem from differences in working memory capacity not already captured by our correction for unaided task performance, we furthermore extended the main model by also including the d’ from a 2-back task as a covariate (H5d). Finally, we fit an extended version of the main model with scores from the cognitive ability test (ICAR5) as an additional covariate to capture cognitive ability (H5e). We ran the same analysis but for the AD factor. We included this test as a preregistered analysis but did not specify any directional hypotheses.

Our final hypothesis, H7, aimed to differentiate between the Metacognitive Monitoring Mechanism, the Metacognitive Control Mechanism, and the Direct Mechanism. We tested how compulsivity would affect the relationship between confidence and offloading. More specifically, we predicted that CIT scores would act as a moderator variable between the metacognitive and the reminder bias and that highly compulsive individuals would have a weaker link. We tested this by fitting the following regression model to the data: biasrem~biasmeta*CIT+AD+age+gender+education+ε

To avoid circularity, we used the unconfounded metacognitive bias and reminder bias for this analysis. The moderation of CIT is reflected in its interaction term with the biasmeta predictor. A significant interaction term can be interpreted as support for the Metacognitive Control Mechanism. In addition to this main model (H7a), we furthermore also tested whether this effect would persist if working memory (2-back d’; H7b) or educational attainment (H7c) were included as additional covariates. We ran the same analysis but for the AD factor. We included this test as a preregistered analysis but did not specify any directional hypotheses.

It should be noted that whilst not explicitly preregistered, our planned models also allow testing for a mediation effect (metacognitive bias acting as a mediator on the effect of the CIT score on the reminder bias). This is done by comparing the effect of CIT on the reminder bias when the effect of the metacognitive bias is accounted for (Hypothesis 7) to when it is not (Hypothesis 5). In addition, we included a causal mediation analysis (not preregistered), using the mediation package in R. This analysis involved testing of the indirect effect using bootstrapping. More specifically, we computed unstandardised indirect effects for each of our 1000 bootstrapped samples and based on those the 95% confidence interval. To keep the information entering into the mediation analysis constant, we re-fitted the models from our sections on H5 and H6/H8 but with the unconfounded metacognitive bias and reminder bias respectively. Furthermore, we had to treat the covariate ‘gender’ as a continuous variable as the mediation package would otherwise not have been able to fit the data. We expect that this difference is unlikely to cause any issues with the interpretation of our effects. A significant mediation effect can be interpreted as support for the Metacognitive Monitoring Mechanism. A significant direct effect can be interpreted as support for the Direct Mechanism.